---
title: qweqweqqwea
description: >-
  This is a template chapter.



---
## <<<New Exercise>>>

```yaml
type: NormalExercise
key: a3a0bc16c1
lang: r
xp: 100
skills: 1
```


`@instructions`

`@hint`

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
---