---
title: qweqweqqwe
description: >-
  This is a template chapter.


---
## popopo

```yaml
type: NormalExercise

xp: 100

key: d04fe714eb
```

test

`@instructions`
test

`@hint`


`@pre_exercise_code`
```{r}
test = 1
```

`@solution`
```{r}
test
```
`@sct`
```{r}
test
```





---
## Insert exercise title here

```yaml
type: NormalExercise

xp: 100

key: cdf5040504
```

asd


`@hint`
asd

`@pre_exercise_code`
```{r}
asd
```
`@sample_code`
```{r}
asd
```
`@solution`
```{r}
asd
```
`@sct`
```{r}
asd
```



