---
title: Template Chapter 1
description: >-
  This is a template chapter.


---
## Insert exercise title here

```yaml
type: NormalExercise

xp: 100

key: d04fe714eb
```

test

`@instructions`
test

`@hint`
test

`@pre_exercise_code`
```{r}
test = 1
```

`@solution`
```{r}
test
```
`@sct`
```{r}
test
```



